Quick-quick-slow: the foxtrot migration and dynamic non-breeding range

Author response:

The following is the authors’ response to the current reviews.

We sincerely appreciate the reviewer’s dedication to evaluating our manuscript and raising essential considerations regarding the classification of the migration behavior we described. While the reviewer suggests that this behavior aligns with the concept of itinerancy, we contend that it represents a distinct phenomenon, albeit with similarities, as both involve the non-breeding movements of birds. We acknowledge that our manuscript did not adequately address this distinction and have considered the reviewer’s feedback. In our response, we clarify the difference between the described phenomenon and itinerancy. Our revised manuscript will include a new section in the Discussion to address this issue comprehensively.

In the first part of the review, the reviewer emphasizes that the pattern we are describing is consistent with itinerancy. Regardless of the terminology used, we want to highlight the existence of two different types of migratory behavior, both of which involve movement in non-breeding areas.

The first type, called itinerancy, was first described by Moreau in 1972 in “The Palaearctic-African Bird Migration Systems.” As noted by the reviewer, this behavior involves an alternation of stopovers and movements between different short-term non-breeding residency areas. They usually occur in response to food scarcity in one part of the non-breeding range, causing birds to move to another part of the same range. These movements typically cover distances of 10 to 100 kilometers but are neither continuous nor directional. Moreau (1972) defined itinerancy as prolonged stopovers, normally lasting several months, primarily in tropical regions. He noted observations of certain species disappearing from his study areas in sub-Saharan Africa in December and others appearing, suggesting they may have multiple home ranges during the non-breeding season. Subsequent research, as mentioned by the reviewer, has confirmed itinerancy in many species, particularly among Palaearctic-African migrants in sub-Saharan Africa. In particular, the Montagu’s Harrier has been extensively studied in this regard. The reviewer rightly points out that our study does not include recent findings on this species. In our revised version, we will include references to recent studies, such as those by Trierweiler et al. (2013, Journal of Animal Ecology, 82:107-120) and Schlaich et al. (2023, Ardea, 111:321-342), which show that Montagu’s Harrier has an average of 3-4 home ranges separated by approximately 200 kilometers. These studies suggest that the species spends approximately 1.5 months at each site, with the most extended period typically observed at the last site before migrating to the breeding grounds.

In the second type, birds undertake a post-breeding migration, arrive in their non-breeding range, and then gradually move in a particular direction throughout the season. This continuous directional movement covers considerable distances and continues throughout the non-breeding period. In our study, this movement covered about 1000 km, comparable to the total migration distance of Rough-legged Buzzards of about 1500 km. As observed in our research, these movements are influenced by external factors such as snow cover. In such cases, the progression of snow cover in a south-westerly direction during winter can prevent birds from finding food, forcing them to continue migrating in the same direction. In essence, this movement represents a prolonged phase of the migration process but at a slower pace. Similar behavior has been documented in buzzards, as reported by Strandberg et al. (2009, Ibis 151:200-206). Although several transmitters in their study stopped working in mid-winter, the authors observed a phenomenon they termed ‘prolonged autumn migration.’

In the second part of the review, the reviewer questions the need to distinguish between the two behaviors we have discussed. However, we believe these behaviors differ in their structure (with the first being intermittent and often non-directional, whereas the second is continuous and directional) and in their causes (with the first being driven by seasonal food resource cycles and the second by advancing snow cover). We therefore argue that it is worth distinguishing between them. To differentiate these forms of non-breeding movement, we propose to use ‘itinerancy’ for the first type, as described initially by Moreau in 1972, and introduce a separate term for the second behavior. Although ‘slow directional itinerancy’ could be considered, we find it too cumbersome.

Moreover, ‘itinerancy’ in the literature refers not only to non-breeding movements but also to the use of different nesting sites, e.g., Lislevand et al. (2020, Journal of Avian Biology: e02595), reinforcing its association with movements between multiple sites within habitats. We, therefore, propose that the second behavior be given a distinct name. We acknowledge the reviewer’s point that we did not adequately address this distinction in the Discussion and plan to include a separate section in our paper’s revised version. In the third part of his review, the reviewer suggests an alternative title. Another reviewer, Dr Theunis Piersma, suggested the current title during the first round of reviewing, and we have chosen his version.

In the fourth part of the review, the reviewer questions whether it is appropriate to discuss the conservation aspect of this study. This type of non-breeding movement raises concerns about accurately determining non-breeding ranges and population dynamics for species that exhibit this behavior. We believe that accurate determination of range and population dynamics is critical to conservation efforts. While this may be less important for species breeding in Europe and migrating to Africa, for which monitoring breeding territories is more feasible, it’s essential for Arctic and sub-Arctic breeding species. Large-scale surveys in these regions have historically been challenging and have become even more so with the end of Arctic cooperation following Russia’s war with Ukraine (Koivurova, Shibata, 2023). For North America and Europe, non-breeding abundance is typically estimated once per season in mid-winter. In North America, these are the so-called Christmas counts (which take place once at the end of December), and in Europe, they are the IWC counts mentioned by the reviewer (as follows from their official website - “The IWC requires a single count at each site, which should be repeated each year. The exact dates vary slightly from region to region, but take place in January or February”). Because of such a single count in mid-winter, non-breeding habitats occupied in autumn and spring will be listed as ‘uncommon’ at best, while south-western habitats where birds are only present in mid-winter will be listed as ‘common.’ However, the situation will be reversed if we consider the time birds spend in these habitats.

The reviewer also highlights the introduction’s unconventional structure and information redundancy at the beginning. We have chosen this structure and provided basic explanations to improve readability for a wider audience, given eLife’s readership. At the same time, we will certainly take the reviewers’ feedback into account in the revised version. We plan to include the references to modern itinerancy research mentioned above and to add a section on itinerancy to the Discussion.

We appreciate the reviewer’s input and sincerely thank them for their time and effort in reviewing our paper. While we may not fully agree on the classification of the behavior we describe, we value the opportunity to engage in discussion and believe that presenting arguments and counterarguments to the reader is beneficial to scientific progress.

The following is the authors’ response to the original reviews.

Reviewer #1 (Recommendations For The Authors):

I much enjoyed reading this manuscript, that is, once I understood what it is about. Titles like "Conserving bird populations in the Anthropocene: the significance of non-breeding movements" are a claim to so-called relevance, they have NOTHING to do with the content of the paper, so once I understood that this paper was about the "Quick quick slow: the foxtrot migration of rough-legged buzzards is a response to habitat and snow" (an alternative title), it was becoming very interesting. So the start of the abstract as well as the introduction is very tedious, as clearly much trouble is taken here to establish reputability. In my eyes this is unnecessary: eLife should be interested in publishing such a wonderful description of such a wonderful migrant in a study that comes to grips with limiting factors on a continental scale!

We sincerely appreciate your time and effort in reviewing our manuscript. Thank you for your appreciation of our study.

We agree that the focus of the article should be changed from conservation to migration patterns. We have rewritten the Introduction and Discussion as suggested. We have added the application of this pattern including conservation at the end of the Discussion by completely changing Figure 5. We have also changed the title to the suggested one.

Not sure that the first paragraph statements that seek to downplay what we know about wintering vs breeding areas are valid (although I see what purpose they serve). Migratory shorebirds have extensively been studied in the nonbreeding areas, for example, including movement aspects (see, as just one example, Verhoeven, M.A., Loonstra, A.H.J., McBride, A.D., Both, C., Senner, N.R. & Piersma, T. (2020) Migration route, stopping sites, and non breeding destinations of adult Black tailed Godwits breeding in southwest Fryslân, The Netherlands. Journal of Ornithology 162, 61-76) and there are very impressive studies on the winter biology of migrants across large scale (for example in Zwarts' Living on the Edge book on the Sahel wetlands). Think also about geese and swans and about seabirds!

We have rewritten the first paragraph and it now talks about patterns of migratory behavior. We have also rewritten the second paragraph, now it is devoted to studies of movements in the non-breeding period. We explain how our pattern differs from those already studied and give references to the papers you mentioned.

Directional movements in nonbreeding areas as a function of food (in this case locusts) have really beautifully been described by Almut Schlaich et al in JAnimEcol for Montagu's harriers.

We have added Montagu's harrier example in the second paragraph of the Introduction and the Discussion. We have added a reference to Schlaich and to Garcia and Arroyo, who suggested that Montagu's harriers have long directional migrations during the non-breeding period.

Once the paper starts talking buzzards, and the analyses of the wonderful data, all is fine. It is a very competent analysis with a description of a cool pattern.

Thank you for your appreciation of our study. We hope the revised version is better and clearer.

However, i would say that it is all a question of spatial scale. The buzzards here respond to changes in food availability, but there is not an animal that doesn't. The question is how far they have to move for an adequate response: in some birds movements of 100s of meters may be enough, and then anything to the scale of rough-legged buzzards.

In the new version of the manuscript, we emphasize that this is a large distance (about 1000 km), comparable to the distance of the fall and spring migrations (about 1400 km) in lines 70-72 of the Introduction and 379-383 of the Discussion.

And actually, several of the shorebirds I know best also do a foxtrot, such as red knots and bar-tailed godwits moulting in the Wadden Sea, then spending a few months in the UK estuaries, before returning to the Wadden Sea before the long migrations to Arctic breeding grounds. The publication of the rough-legged buzzard story may help researchers to summarize patterns such as this too. Mu problem with this paper is the framing. A story on the how and why of these continental movements in response to snow and other habitat features would be a grand contribution. Drop Anthropocene, and rethink whether foxtrot should be introduced as a hypothesis or a summary of cool descriptions. I prefer the latter, and recommend eLife to go with that too, rather than encourage "disconnected frames that seek 'respectability'" Good luck, theunis piersma

We thank the reviewer again for his valuable comments and suggestions. We have changed the framing to the suggested one and removed the Anthropocene from the article.

Reviewer #2 (Recommendations For The Authors):

We sincerely appreciate the time and effort you have taken to review our manuscript. We have carefully considered all of your comments, including both public and author comments, and provided detailed responses to each of them below. In addition, we would like to address the most important public comments.

We agree with the suggestion to shift the focus of the article from conservation to migration patterns. Accordingly, we have rewritten both the Introduction and Discussion sections to focus on migration behavior rather than conservation.

However, we respectfully disagree with the suggestion that the migration patterns we describe are synonymous with itinerancy. We acknowledge that our original presentation may have been unclear and may have hindered full understanding. In the revised version, we provide a detailed analysis of migratory behavior in the Introduction that describes how our pattern differs from itinerancy. We also revisit this distinction in the Discussion section. We have also carefully revised Figure 1 to improve clarity and avoid potential misunderstandings.

Regarding the applicability of the described migration pattern, we acknowledge that the Rough-legged Buzzard is not listed as an endangered species. However, we believe that our findings have practical implications. We have moved our discussion of this issue to the end of the Discussion section and have completely revised Figure 5. While the overall population of Rough-legged Buzzards is not declining, certain regions within its range are experiencing declines. We show that this decline does not warrant listing the species as endangered. Instead, it may represent a redistribution within the non-breeding range - a shift in range dynamics. We use the example of the Rough-legged Buzzard to illustrate this concept and emphasize the importance of considering such dynamics when assessing the conservation status of species in the future.

We also acknowledge that the hypothesis of this form of behavior has been proposed previously for Montagu's Harrier, and we have included this information in the revised manuscript. In addition, we agree that the focus on the Anthropocene is unnecessary in this context and have therefore removed it.

We believe that these revisions significantly improve the clarity and robustness of the manuscript, and we are grateful for your insightful comments and suggestions.

As a general comment, please note that including line numbers (as it is the standard in any manuscript submission) would facilitate reviewers providing more detailed comments on the text.

We apologize for this oversight and have added line numbers to our revised manuscript.

Dataset: unclear what is the frequency of GPS transmissions. Furthermore, information on relative tag mass for the tracked individuals should be reported.

We have included this information in our manuscript (L 157-163). We also refer to the study in which this dataset was first used and described in detail (L 164).

Data pre-processing: more details are needed here. What data have been removed if the bird died? The entire track of the individual? Only the data classified in the last section of the track? The section also reports on an 'iterative procedure' for annotating tracks, which is only vaguely described. A piecewise regression is mentioned, but no details are provided, not even on what is the dependent variable (I assume it should be latitude?).

Regarding the deaths. We only removed the data when the bird was already dead. We have corrected the text to make this clear (L 170).

Regarding the iterative procedure. We have added a detailed description on lines 175-188.

Data analysis: several potential issues here:

(1) Unclear why sex was not included in all mixed models. I think it should be included.

Our dataset contains 35 females and eight males. This ratio does not allow us to include sex in all models and adequately assess the influence of this factor. At the same time, because adult females disperse farther than males in some raptor species, we conducted a separate analysis of the dependence of migration distance on sex (Table S8) and found no evidence for this in our species. We have written a separate paragraph about this. This paragraph can be found on lines 356-360 of the new manuscript.

(2) Unclear what is the rationale of describing habitat use during migration; is it only to show that it is a largely unsuitable habitat for the species? But is a formal analysis required then? Wouldn't be enough to simply describe this?

Habitat use and snow cover determine the two main phases (quick and slow) of the pattern we describe. We believe that habitat analysis is appropriate in this case and that a simple description would be uninformative and would not support our conclusions.

(3) Analysis of snow cover: such a 'what if' analysis is fine but it seems to be a rather indirect assessment of the effect of snow cover on movement patterns. Can a more direct test be envisaged relating e.g. daily movement patterns to concomitant snow cover? This should be rather straightforward. The effectiveness of this method rests on among-year differences in snow cover and timing of snowfall. A further possibility would be to demonstrate habitat selection within the entire non-breeding home range of an individual in relation snow cover. Such an analysis would imply associating presence-absence of snow to every location within the non-breeding range and testing whether the proportion of locations with snow is lower than the proportion of snow of random locations within the entire non-breeding home range (95% KDE) for every individual (e.g. by setting a 1/10 ratio presence to random locations).

The proposed analysis will provide an opportunity to assess whether the Rough-legged Buzzard selects areas with the lowest snow cover, but will not provide an opportunity to follow the dynamics and will therefore give a misleading overall picture. This is especially true in the spring months. In March-April, Rough-legged Buzzards move northeast and are in an area that is not the most open to snow. At this time, areas to the southwest are more open to snow (this can be seen in Figure 4b). If we perform the proposed analysis, the control points for this period would be both to the north (where there is more snow) and to the south (where there is less snow) from the real locations, and the result would be that there is no difference in snow cover.

A step-selection analysis could be used, as we did in our previous work (Curk et al 2020 Sci Rep) with the same Rough-legged Buzzard (but during migration, not winter). But this would only give us a qualitative idea, not a quantitative one - that Rough-legged Buzzards move from snow (in the fall) and follow snowmelt progression (in the spring).

At the same time, our analysis gives a complete picture of snow cover dynamics in different parts of the non-breeding range. This allows us to see that if Rough-legged Buzzards remained at their fall migration endpoint without moving southwest, they would encounter 14.4% more snow cover (99.5% vs. 85.1%). Although this difference may seem small (14.4%), it holds significance for rodent-hunting birds, distinguishing between complete and patchy snow cover. Simultaneously, if Rough-legged Buzzards immediately flew to the southwest and stayed there throughout winter, they would experience 25.7% less snow cover (57.3% vs. 31.6%). Despite a greater difference than in the first case, it doesn't compel them to adopt this strategy, as it represents the difference between various degrees of landscape openness from snow cover.

We write about this in the new manuscript on lines 385-394.

Results: it is unclear whether the reported dispersion measures are SDs or SEs. Please provide details.

For the date and coordinates of the start and end of the different phases of migration, we specified the mean, sd, and sample size. We wrote this in line 277. For the values of the parameters of the different phases of the migration (duration, distance, speed, and direction), we used the mean, the standard error of the mean, and the confidence interval (obtained using the ‘emmeans’ package). We have indicated this in lines 302-303 and the caption of Table 1 (L 315) and Figure 2 (L 293-294). For the values of habitat and snow cover experienced by the Rough-legged Buzzards, we used the mean and the error of the mean. We reported this on lines 322 and 337 and in Figures 3 (L 332-333) and 4 (L 355-356).

Discussion: in general, it should be reshaped taking into account the comments. It is overlong, speculative and quite naive in several passages. Entire sections can be safely removed (I think it can be reduced by half without any loss of information). I provide some examples of the issues I have spotted below. For instance, the entire paragraph starting with 'Understanding....' is not clear to me. What do you mean by 'prohibited management' options? Without examples, this seems a rather general text, based on unclear premises when related to the specific of this study. Some statements are vague, derive from unsubstantiated claims, and unclear. E.g. "Despite their scarcity in these habitats, forests appear to hold significant importance for Rough-legged buzzards for nocturnal safety". I could not find any day-night analysis showing that they actually roost in forests during nighttime. Being a tundra species, it may well be possible that rough-legged buzzards perceive forests as very dangerous habitats and that they prefer instead to roost in open habitats. Analysing habitat use during day and night during the non-breeding period may be of help to clarify this. Furthermore, considering the fast migration periods, what is the flight speed during day and night above forests? Do these birds also migrate at night or do they roost during the night? Perhaps a figure visualizing day and night track segments could be of help (or an analysis of day vs. night flight speed) (there are several R packages to annotate tracks in relation to day and night). This is an example of another problematic statement: "The progression of snow cover in the wintering range of Rough-legged buzzards plays a significant role in their winter migration pattern." The manuscript does not contain any clear demonstration of this, as I wrote in my previous comments. Without such evidence, you must considerably tone down such assertions. But since providing a direct link is certainly possible, I think that additional analyses would clearly strengthen your take-home message.

The paragraph starting with "The quantification of environmental changes that could prove fatal to bird species presents yet another challenge for conservation efforts in an era of rapid global change." is quite odd. Take the following statement "For instance, the presence of small patches of woodland in the winter range might appear crucial to the survival of the Rough-legged buzzard. Elimination of these seemingly minor elements of vegetation cover through management actions could have dire consequences for the species.". It is based on the assumption that minor vegetation elements play a key role in the ecology of the species, without any evidence supporting this. Does it have any sense? I could safely say exactly the opposite and I would believe it might even be more substantiated.

We agree with these comments.

We have completely rewritten this section. As suggested, we have shortened it by removing statements that were not supported by the research. We have completely removed the statements about "prohibited management". We have also removed the statement that "forests appear to be of significant importance to Rough-legged buzzards for nocturnal safety" and everything associated with that statement, e.g. the statement about "small elements of vegetation cover", etc. We do believe that this statement is true in substance, but we also agree that it is not supported by the results and requires separate analysis. At the same time, we believe that this is a topic for a separate study and would be redundant here. Therefore, we leave it for a separate publication.

Conclusion paragraph: I believe this severely overstates the conservation importance of this study. That the results have "crucial implications for conservation efforts in the Anthropocene, where rapidly changing environmental factors can severely impact bird migration" seems completely untenable to me. What is the evidence for such crucial implications? For instance, these results may suggest that climate change, because global warming is predicted to reduce snow cover in the non-breeding areas, might well be beneficial for populations of this species, by reducing non-breeding energy expenditure and improving non-breeding survival. I think statements like these are simply not necessary, and that the study should be more focused on the actual results and evidence provided.

We have completely rewritten this section. We removed the reference to the Anthropocene and focused on migratory behavior and migration patterns.

Changes from the earlier version to the revised version

Abstract

Abstract

Added in the revised version: Advancements in tracking technologies have revolutionized our understanding of bird migrations, revealing a diverse array of migratory behaviors. We propose a new pattern of migratory behavior termed ‘foxtrot migration,’ characterized by alternating quick and slow phases during the non-breeding period. This behavior involves directional and continuous movements, distinct from traditional seasonal itinerancy or forage-on-fly migration. Using the Rough-legged buzzard as a model species, we confirmed the presence of foxtrot migration and ‘dynamic non-breeding ranges’, driven by environmental factors such as snow cover dynamics. We advocate for accurate representation of dynamic ranges on maps and emphasize the need to consider range dynamics when assessing species conservation status. Our findings underscore the importance of understanding complex migratory behaviors in the face of environmental change, facilitated by advancements in tracking technologies. This knowledge is crucial for effective conservation strategies amid ongoing global environmental challenges.

Removed in the revised version: Background

Removed in the revised version: To conserve bird species threatened by climate change, it is important to understand how environmental factors affected by climate change, such as snow cover, impact their ranges. While this problem is fairly well understood for breeding areas, it remains poorly understood for non-breeding areas. In non-breeding areas, seasonal cycles can strongly influence the distribution of resources during winter. If birds adapt to such changes, they may result in seasonal and directional movement of birds within their non-breeding range. In this case, birds would experience a unique migration pattern - rapid migration between breeding and non-breeding habitats versus a slow migration pattern within their non-breeding range. Their non-breeding range would therefore be dynamic, with potentially important consequences for our understanding of population densities and non-breeding ranges.

Removed in the revised version: Methods

Removed in the revised version: Between 2013-2021, we tracked 43 adult Rough-legged buzzards with solar GPS-GSM loggers. We analyzed their behavior, determined whether the birds showed any directional return migrations during the non-breeding season, and evaluated the differences between the slow migration within their winter range and the quick migration between breeding and non-breeding areas. We also analyzed the vegetation cover of the areas crossed during quick and slow migrations and the role of snow cover in winter migrations.

Removed in the revised version: Results

Removed in the revised version: Our findings revealed that after a quick fall migration through the taiga zone, Rough-legged buzzards continue to migrate during the non-breeding season, albeit at a slower pace across the wooded fields they select as habitat. They avoid complete snow cover and move to escape the progression of the snow cover line from northeast to southwest and back during the winter. As a consequence, Rough-legged buzzards have a dynamic winter range. Thus, the migration pattern of these birds comprises alternating quick and slow phases, resembling the foxtrot dance, which we have named the ‘foxtrot’ migration pattern. Due to this pattern, their winter range displays a dynamic shift of the seasonal center of the population distribution over 1000 km towards the southwest and back throughout the winter.

Removed in the revised version: Conclusions

Removed in the revised version: Our study uncovered a novel bird migration pattern postulated to exist before but poorly understood. This ‘foxtrot migration’ likely occurs in many migratory species inhabiting winter areas with pronounced seasonal cycles. Our findings have implications for conservation efforts in the Anthropocene, where environmental factors such as snow cover can change rapidly and have cascading effects on bird migration. We recommend presenting dynamic winter ranges in species descriptions and range maps so ecologists can use them to develop effective conservation strategies.

Introduction

Added in the revised version: In recent decades, advanced tracking technology has greatly improved our knowledge of bird migrations (Wikelski et al. 2007; Jetz et al. 2022; Kays & Wikelski 2023). In the 19th century, the finding of a White stork (Ciconia ciconia) in Germany with an African arrow led to the discovery that birds did not hibernate at the bottom of water bodies during the winter but rather made extensive seasonal migrations (Richter & Bick 2018). Subsequent studies of bird migration in the 20th century using ringing techniques revealed a diversity of seasonal migratory behaviors, including 1) propensity to migrate (migrants, residents, partial migrants), 2) differences among populations (leapfrog, chain, parallel migrations), 3) connectivity and philopatry (constancy of breeding and non-breeding locations), 4) length of migration routes (long-distance and short-distance migrants), 5) variations in flight paces during migration (with stopovers, forage-on-fly, nonstop), and 6) tendencies to move within a single season (breeding and non-breeding itinerancy) (Alerstam, Hedenström & Åkesson 2003; Newton 2008; Alerstam 2011; Berthold, Gwinner & Sonnenschein 2013). Different combinations of these behaviors contribute to a wide variety of bird migration patterns (Chapman et al. 2014; Alerstam & Bäckman 2018; Lislevand et al. 2020). However, the advent of 21st-century technology, such as GPS transmitters, provides a higher resolution of movement and the ability to observe more details of migratory behavior (Wikelski et al. 2007; Jetz et al. 2022). Just as advances in image resolution allow us to see finer details in images, improvements in bird tracking technology now allow us to identify previously undiscovered migratory behaviors, revealing new patterns.

Added in the revised version: Current technologies allow observing forms of migratory behavior that are intermediate or integrative to those described above. Itinerancy was initially defined as the movements of Afro-Palearctic migrants during the non-breeding season in the Sahel, with a combination of stopovers and flights in between (Moreau 1972). Subsequently, researchers investigated this behavior for multiple species wintering in Africa, attributing their movements during the non-breeding season to changes in food availability predicted by NDVI (Zwarts et al. 2012; Schlaich et al. 2016; Thorup et al. 2017; Verhoeven et al. 2021). In 1998, Garcia and Arroyo hypothesized that certain species, like the Montagu’s harrier (Circus pygargus), might not only cover tens to hundreds of kilometers between various sites in their winter range but could also demonstrate directional, continuous movement throughout the entire non-breeding season. This movement may span a distance similar to their migration distance (ca. 1000 km), essentially indicating a prolonged, slower-paced continuation of migration (García & Arroyo 1998). Because this behavior is intermediate between seasonal movement (itinerancy) and the actual migration process, we can describe it as either ‘persistent directional non-breeding itinerancy’ or ‘very slow forage-on-fly migration in the non-breeding area’. Both options are clumsy and can confuse. We suggest naming this behavior with a new term to avoid confusion with established terms such as ‘itinerancy’ for any seasonal movement and ‘forage-on-fly’ denoting movement from breeding to non-breeding sites (Alerstam, Hake & Kjellén 2006; Alerstam 2011; Lopez-Ricaurte et al. 2024). Considering this behavior as an extension of migration and looking at the complete life cycle of such species, it includes a quick phase during the transition between breeding and non-breeding areas, followed by a slow phase of directed and seasonal movement within the non-breeding area, and finally another quick phase towards the breeding area. Analogous to the alternating fast and slow movements of the foxtrot dance, we propose the term ‘foxtrot migration’ for this seasonal movement to provide a concise and easily understandable description (Figure 1a). Consequently, we suggest to term the non-breeding range of species exhibiting this migration type as ‘dynamic range’ (Figure 1b).

Removed in the revised version: Global climate change significantly impacts organisms and regions worldwide, with the Arctic experiencing more pronounced effects (Davy & Outten 2020; Previdi, Smith & Polvani 2021). These changes are potentially responsible for the decline in the numbers and ranges of many species and populations (Post et al. 2009; Gilg et al. 2012; Schmidt et al. 2012). Understanding the impact of environmental factors highly susceptible to climate change variation, such as snow cover, is critical for species conservation planning (Lawler 2009; Groves et al. 2012; Reside, Butt & Adams 2018). However, determining the distribution range and the seasonal changes in the local density of highly mobile taxa such as birds can be challenging. Many bird species undergo seasonal migrations, traditionally dividing their range into breeding and non-breeding areas. Most studies of climate change’s impact on bird species have focused on changes in the breeding range, with less emphasis on the non-breeding range (Schmidt et al. 2012). This is primarily because the breeding range is easy to determine from nest counts and does not change during the season. In contrast, the non-breeding range is difficult to define as birds are not tied to any one location, despite spending most of their annual cycle within this range. Moreover, environmental factors affected by climate change can significantly influence the dynamics of the non-breeding range. In recent decades, advanced tracking technology has greatly improved our ability to study bird ecology. We can now follow birds throughout their life cycle and study the influence of environmental factors on the dynamic of their non-breeding range (Wikelski et al. 2007; Flack et al. 2022; Jetz et al. 2022).

Removed in the revised version: The non-breeding range of many bird species can experience dramatic seasonal changes in food availability due to environmental conditions, such as the progressive snow cover line (Bildstein 2006). This line moves – at least in most north-temperate areas exposed to the northwesterly wind systems – from north (-east) to south (-west) and back again in many mid-latitude areas between October and May, affecting food availability for many bird species (Sonerud 1986; Vansteelant, Faveyts & Buckens 2011). We propose that species dependent on this environmental factor move gradually away from the snow cover during the winter and then gradually move in the opposite direction, resulting in a directed and seasonal movement of their winter range. Considering the complete life cycle of such species, they would exhibit a quick migration phase when moving between breeding and non-breeding ranges, followed by a slow migration phase with directed and seasonal movement within the non-breeding range, and finally, another quick migration phase towards the breeding range. Many insectivore bird species wintering in Africa exhibit a similar pattern – relatively fast migration across the Sahara desert and then experiencing gradual directional movements during the non-breeding period (Stach et al. 2012; Jacobsen et al. 2017). Researchers propose that this movement is a consequence of a shift in the insect peak, which is, in turn, influenced by seasonal changes in precipitation (Stach et al. 2012). This type of movement can be compared to the foxtrot dance, which also involves alternating quick and slow phases of movement, both forward and backward. We, therefore, refer to this putative migration as the ‘foxtrot’ migration (Figure 1a) to allow for a quick and easily understandable description. The winter range in species exhibiting this type of migration will then have a dynamic structure that shifts with the ongoing non-breeding migration over time (Figure 1b).

Earlier version: Hypothetical foxtrot migration and dynamic winter range scheme.

Revised version: Foxtrot migration and dynamic winter range scheme.

Added in the revised version: a) Foxtrot migration. The color bar along the y-axis corresponds to the color coding of the habitats in panel b of the chart on the right. b) Dynamic winter range.

Added in the revised version: The concept of dynamic range, derived from foxtrot migration, is critical to understanding the current range declines and shifts observed in many species. Typically, a species’ range is divided into areas of common and uncommon, with temporal heterogeneity limited to breeding and non-breeding distinctions (Keller et al. 2020; Billerman et al. 2024). However, suppose a species exhibits a foxtrot migration, moving gradually in one direction and back again during the non-breeding season. In this case, it becomes inaccurate to represent its non-breeding range as static. This discrepancy is particularly pronounced when ranges are mapped using mid-winter surveys, such as Christmas counts. During this period, the species may be present at the outer limits of its non-breeding range, leading to a false classification as common in that area when, in fact, it is rare if the entire non-breeding season is considered. Similarly, extended negative population dynamics of a species in a particular part of its non-breeding range could be misinterpreted as a threat to the entire population. Indeed, if this species has a foxtrot migration, range dynamics may have changed, possibly resulting in the species ceasing to migrate to that particular part of its range. Advanced technologies now make it possible the identification of foxtrot migration and range dynamics. This, in turn, facilitates accurate range mapping and proper assessment of threats to the species.

Added in the revised version: This pattern of behavior has been proposed in species that overwinter in Africa, where non-breeding movements are primarily related to the distribution of insect abundance (García & Arroyo 1998). At the same time, it is reasonable to assume that this behavior may also be characteristic of species that spend the non-breeding season in mid-latitudes, where food availability may be linked to snow cover dynamics. These dynamics have changed particularly rapidly in recent decades due to global warming (Davy & Outten 2020; Previdi, Smith & Polvani 2021). As a result, global climate change may significantly affect birds that breed in the Arctic and winter in mid-latitudes. The Arctic is experiencing more pronounced effects of global change than other regions, with declines in the abundance and range of several species and populations (Post et al. 2009; Gilg et al. 2012; Schmidt et al. 2012). Therefore, understanding how snow cover dynamics affect Arctic species across habitats throughout the annual cycle is essential today. Furthermore, elucidating the foxtrot migration and dynamic non-breeding range of mid-latitude wintering Arctic birds may provide valuable information on expected future changes in their non-breeding range.

Removed in the revised version: a) Foxtrot migration, Q. mig. – Quick migration (spring and fall migrations). b) Dynamic winter range.

Removed in the revised version: This concept of ‘foxtrot’ migration synthesizes migrating and non-breeding movements. Traditionally, studies of bird movement ecology fall into two broad categories: migration studies and studies of movements within breeding and non-breeding ranges (Newton 2008; Berthold, Gwinner & Sonnenschein 2013). Migration studies have focused on the directed seasonal return movements between breeding and non-breeding ranges (Alerstam, Hedenström & Åkesson 2003; Bildstein 2006; Curk et al. 2020). This includes studies of different aspects of migration behavior: reasons for migration (McKinnon et al. 2010; Somveille, Rodrigues & Manica 2018; Pokrovsky et al. 2021), the timing of migration (Lameris et al. 2018; Lehikoinen et al. 2019), duration and length of migration routes (Nilsson et al. 2013), and other aspects. The variability in these characteristics gives rise to a wide range of migratory patterns, which vary greatly in the time taken to migrate. Some species migrate non-stop to their wintering grounds (Battley et al. 2012), some species experience a ‘fly-and-forage’ strategy, where birds hunt during migration and reduce their stopover time (Strandberg and Alerstam 2007), some make long stops to refuel (Shimada et al. 2014). Research in this area has also examined the influence of external environmental factors such as wind, rain, temperature, and day length on bird migration, revealing how these factors can limit food availability and shape migration patterns (Liechti & Bruderer 1998; Alerstam, Hedenström & Åkesson 2003; Curk et al. 2020; Pokrovsky et al. 2021). The second area of research focuses on daily routine movements, dispersal movements, nomadism, and habitat selection within breeding and non-breeding habitats (Greenwood & Harvey 1982; Lenz et al. 2015; Curk et al. 2022). In these studies, it is common to observe long movements in time and distance that are non-directional and non-returning as birds search for food or suitable nesting sites. At the same time, seasonal cycles can significantly affect resource distribution across the wintering range of species. This, in turn, can lead to directional and seasonal bird movements within the wintering range, commonly referred to as migratory behavior. We hypothesize that such a migration pattern and dynamic winter range could be observed in mid-latitude wintering Arctic raptors.

Earlier version: We used the Rough-legged buzzard (Buteo lagopus) as a model species to investigate this phenomenon. The Rough-legged buzzard is an Arctic breeding and mid-latitude wintering raptor (Ferguson-Lees & Christie 2001; Bechard & Swem 2002). Rough-legged buzzards feed mainly on small rodents in the Arctic and its wintering grounds (Tast, Kaikusalo & Lagerström 2010; Pokrovsky et al. 2014). They prefer open areas for hunting, and trees and tall bushes or uplands for resting. In the Arctic, such areas are the southern and typical tundra (Walker et al. 2005), and in the mid-latitudes, areas with fields and patches of forests (wooded fields). The taiga zone, where there is little open space, is unsuitable for them, although they may nest in the northern taiga zone on the border with the tundra zone (Sundell et al. 2004). Snow cover and day length play an important role in their life (Terraube et al. 2015; Curk et al. 2020; Pokrovsky et al. 2021). Rough-legged buzzards can only hunt during the day (Pokrovsky et al. 2021), and heavy snow cover makes hunting for small rodents problematic (Sonerud 1986; Vansteelant, Faveyts & Buckens 2011). These two environmental factors, therefore, affect the availability of prey for Rough-legged buzzards. At the same time, these factors vary considerably in mid-latitudes during winter. Thus, prey availability for Rough-legged buzzards in the mid-latitudes increases until mid-winter if they migrate southwards and after mid-winter if they migrate northwards. We, therefore, assume that Rough-legged buzzards could track prey availability and experience a ‘foxtrot’ migration and dynamic wintering distribution.

Revised version: We used the Rough-legged buzzard (Buteo lagopus) as a model species to investigate this phenomenon. The Rough-legged buzzard is an Arctic breeding and mid-latitude wintering raptor (Ferguson-Lees & Christie 2001; Bechard & Swem 2002). Rough-legged buzzards feed mainly on small rodents in the Arctic and its wintering grounds (Tast, Kaikusalo & Lagerström 2010; Pokrovsky et al. 2014). They prefer open areas for hunting, and trees and tall bushes or uplands for resting. In the Arctic, such areas are the southern and typical tundra (Walker et al. 2005), and in the mid-latitudes, areas with fields and patches of forests (wooded fields). The taiga zone, where there is little open space, is unsuitable for them, although they may nest in the northern taiga zone on the border with the tundra zone (Sundell et al. 2004). Snow cover and day length play an important role in their life (Terraube et al. 2015; Curk et al. 2020; Pokrovsky et al. 2021). Rough-legged buzzards can only hunt during the day (Pokrovsky et al. 2021), and heavy snow cover makes hunting for small rodents problematic (Sonerud 1986; Vansteelant, Faveyts & Buckens 2011). These two environmental factors, therefore, affect the availability of prey for Rough-legged buzzards. At the same time, these factors vary considerably in mid-latitudes during winter. Thus, prey availability for Rough-legged buzzards in the mid-latitudes increases until mid-winter if they migrate southwards and after mid-winter if they migrate northwards. We, therefore, assume that Rough-legged buzzards could track prey availability and experience a foxtrot migration and dynamic non-breeding range.

Earlier version: For this study, we hypothesized that Rough-legged buzzards would undertake a combined ‘foxtrot’ migration and dynamic winter range in response to seasonal changes in mid-latitude environmental factors. We made the following predictions. 1) Buzzards would exhibit a directional and seasonal movement pattern during the winter, moving from the northeast to the southwest and back again. This would result in a dynamic winter range that would continue to move geographically throughout the season. 2) Winter migrations would differ from fall and spring migrations in duration, extent, speed, and direction. 3) Winter migrations would occur in suitable open habitats, whereas fall and spring migrations would occur in unfitting forested areas. 4) The snow conditions experienced by the Rough-legged buzzards during their winter migration would differ from those in the areas of their winter range where they arrived after their autumn migration, with less snow cover encountered during their winter migration than if they had remained where they arrived at the end of their fall migration.

Revised version: For this study, we hypothesized that Rough-legged buzzards undertake a combined foxtrot migration and dynamic non-breeding range in response to seasonal changes in mid-latitude environmental factors. We made the following predictions. 1) Buzzards would exhibit a directional and seasonal movement pattern during the non-breeding period, moving from the northeast to the southwest and back again. This would result in a dynamic non-breeding range that would continue to move geographically throughout the season. 2) Non-breeding movements would occur in suitable open habitats, whereas fall and spring migrations would occur in unfitting forested areas. 3) Non-breeding movements would differ from fall and spring migrations in duration, extent, speed, and direction, however would continue throughout the season. 4) During non-breeding movements, Rough-legged buzzards will experience less snow cover than if they had stayed where they arrived at the end of the fall migration.

Earlier version: In the following, we will refer to the spring and fall migrations as the quick migration, the winter migration to the lowest point of latitude as the 1st phase of the slow migration, and the migration from the lowest point of latitude to the starting point of the spring migration as the 2nd phase of the slow migration (Figure 1a).

Revised version: In the following, we will refer to the spring and fall migrations as the quick phase, the non-breeding movement to the lowest point of latitude as the 1st part of the slow phase, and the movement from the lowest point of latitude to the starting point of the spring migration as the 2nd part of the slow phase (Figure 1).

Material and methods

Dataset

For this study, we tracked 43 adult Rough-legged buzzards (35 females and eight males) with the solar GPS-GSM loggers (e-obs GmbH and UKn – University of Konstanz). The fieldwork was carried out in the Russian Arctic in 2013-2019 at four study sites: Kolguev Island (69°16′N, 48°87′E), Nenetsky Nature Reserve (68°20′N, 53°18′E), Vaigach Island (69°43′N, 60°08′E), and Yamal Peninsula (68°12′N, 68°59′E). For details on capture methods and tag parameters, see Curk et al. (2022); for detailed study descriptions, see Pokrovsky et al. (2015) for Kolguev and Pokrovsky et al. (2019) for Yamal and Nenetsky.

Earlier version: During data pre-processing, we estimated the date of death using the accelerometer or GPS data and removed the tracking data if the bird was dead. We then removed duplicated timestamps and calculated the mean daily positions of each individual. We partitioned the resulting dataset into several periods: 1) breeding, 2) fall migration, 3)1st phase of winter, 4) 2nd phase of winter, and 5) spring migration. We estimated the migration dates – the start and stop dates of the spring and fall migrations – using an iterative search procedure for piecewise regression described by Crawley (2007). We estimated the date between winter’s first and second phases as the day when the mean daily latitude was minimum.

Revised version: During data pre-processing, we estimated the date of death using the accelerometer or GPS data and removed the tracking data if the bird was dead. We then removed duplicated timestamps and calculated the mean daily positions of each individual. We partitioned the resulting dataset into several periods: 1) breeding, 2) fall migration, 3)1st part of winter, 4) 2nd part of winter, and 5) spring migration. We estimated the migration dates – the start and stop dates of the spring and fall migrations – using an iterative search procedure for piecewise regression described by Crawley (2007). We estimated the date between winter’s first and second parts as the day when the mean daily latitude was minimum.

Data analysis

First, we used linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’ (Bates et al. 2015)) to investigate whether or not Rough-legged buzzards migrated during winter. Latitude was the response variable, day of the year was a fixed effect, and individuals and year were included as random effects. Analyses were conducted separately for each migration period (fall, first phase of winter, second phase of winter, and spring). For both phases of the winter migration, we analyzed two additional models with longitude as the response variable instead of latitude. Likelihood ratio tests were used to compare candidate models. The year was not a calendar year but a year between two consecutive breeding seasons. Thus, fall migration, consecutive winter, and consecutive spring have the same value for the year. The day of the year was recalculated consecutively.

Second, we used linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’) to investigate whether migrations’ parameters differ between the migration periods. We analyzed four migration parameters: distance, duration, speed, and direction. The distance was calculated as the distance between two coordinates (start and end of migration) using the R function ‘distm’ in the library ‘geosphere’ (Hijmans 2016). The duration was calculated as the number of days between the start and end of migration. Speed was calculated as the ratio of distance to duration. The direction was calculated as the bearing from the start of the migration coordinates to the end of the migration coordinates using the R function ‘bearing’ in the library ‘geosphere’ (Hijmans 2016). The migration parameter was used as the response variable, the type of migration as a fixed factor, and individuals as a random factor. Likelihood ratio tests were used to compare candidate models. We considered four different parameters of migration (distance, duration, speed, and direction) and four types of migration (fall, first phase of winter, second phase of winter, and spring). The analysis was done separately for each of the migration parameters. Then, we used post hoc comparisons using the R function ‘emmeans’ in the library ‘emmeans’ (Lenth et al. 2019) to compare the estimated means. In some raptor species, adult females disperse further than males (Mearns & Newton 1984; Serrano et al. 2001; Bildstein 2006; Whitfield et al. 2009). Therefore, we conducted an additional analysis on the effect of sex on migration length using linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’). The migration distance was used as the response variable, sex as a fixed factor, and individuals as a random factor.

Third, we investigated whether vegetation land cover differed between areas crossed during the quick (fall and spring) and slow (winter) migrations. We used the combined Terra and Aqua Moderate Resolution Imaging Spectroradiometer (MODIS) Land Cover Climate Modelling Grid (CMG) (MCD12C1) version 6 dataset (Friedl & Sulla-Menashe 2015). We used a modified Leaf Area Index (LAI) as a classification scheme. We combined all four forest types and savannas into one category (forest) and excluded the categories: water bodies, and unclassified. We, therefore, had five types of vegetation cover: forest, grassland, cropland, shrubland, and urban. We annotated the mean daily positions with the vegetation cover type using the Env-DATA tool (Dodge et al. 2013). We used general linear mixed effects models with a binomial distribution (R function ‘glmer’ in the library ‘lme4’) to investigate whether vegetation cover types differ between migration periods. Presence/absence of the studied vegetation land cover type was used as a response variable, migration type as a fixed factor, and individuals as a random factor. The analysis was done separately for each of the vegetation land cover types.

Removed in the revised version: Fourth, we investigated whether snow cover could drive the slow migration phenomenon.

Earlier version: We then compared the snow cover conditions the birds experienced during the winter with two hypothetical snow cover conditions that the birds would have experienced if they had not migrated during the winter. The first hypothetical snow cover condition would have happened if the birds had stayed where they arrived from the north (i.e., where their fall migration ended). To estimate this parameter, we calculated the winter dynamics of the average snow cover at the minimum convex polygons (MCP) occupied by the birds in October and April (northeast of their winter range). A second hypothetical snow cover condition would be if the birds flew immediately to the southwest and spent the whole winter there. To evaluate this, we calculated the winter dynamics of average snow cover on the MCPs occupied by the birds in January and February (southwest of their winter range). We then compared the values obtained for the real snow cover and two hypothetical snow covers using general linear mixed effects models with a binomial distribution (R function ’glmer’ in the ’lme4’ library). Presence/absence of snow cover was used as a response variable, type of snow cover (real, 1st hypothetical, or 2nd hypothetical) as a fixed factor, and years as a random factor. The analysis was done separately for each month. We then used post hoc comparisons to compare the estimated means, using the R function ‘emmeans’ in the ‘emmeans’ library (Lenth et al. 2019).

Revised version: Fourth, we investigated whether snow cover could drive the slow migration phenomenon. We then compared the snow cover conditions the birds experienced during the winter with two hypothetical snow cover conditions that the birds would have experienced if they had not migrated during the winter. The first hypothetical snow cover condition would have happened if the birds had stayed where they arrived from the north (i.e., where their fall migration ended). To estimate this parameter, we calculated the winter dynamics of the average snow cover at the minimum convex polygons (MCP) occupied by the birds in October and April (northeast of their winter range). A second hypothetical snow cover condition would be if the birds flew immediately to the southwest and spent the whole winter there. To evaluate this, we calculated the winter dynamics of average snow cover on the MCPs occupied by the birds in January and February (southwest of their winter range). We then compared the values obtained for the real snow cover and two hypothetical snow covers using general linear mixed effects models with a binomial distribution (R function ‘glmer’ in the ‘lme4’ library). Presence/absence of snow cover was used as a response variable, type of snow cover (real, 1st hypothetical, or 2nd hypothetical) as a fixed factor, and years as a random factor. The analysis was done separately for each month. We then used post hoc comparisons to compare the estimated means, using the R function ‘emmeans’ in the ‘emmeans’ library (Lenth et al. 2019).

We obtained monthly snow cover data with a spatial resolution of ca 500 meters (Global SnowPack MODIS) from the German Aerospace Center (DLR). This product is based on the Moderate Resolution Imaging Spectroradiometer (MODIS) daily snow cover products MOD10A1 and MYD10A1 (version 6 as provided by the National Snow and Ice Data Center NSIDC), which have been processed to remove the gaps due to cloud cover and polar darkness (Dietz, Kuenzer & Dech 2015). These processing steps include a combination of data available from different satellites (Aqua and Terra), 3-day temporal moving window filtering, a regional snow line elevation interpolation relying on a Digital Elevation Model (DEM), and a seasonal filter running through the time series for the whole hydrological year (1st of September through August 31st). The proportion of days in which one pixel is snow-covered per month is referred to here as fractional snow cover and is derived from these daily gap-filled rasters. Five MODIS tiles (h19v03, h20v03, h20v04, h21v03 and h21v04) were mosaicked and re-projected to WGS84. Then, for each month from October to April, we calculated 95% minimum convex polygons (MCPs) for the distribution of Rough-legged buzzards using the R function ‘mcp’ in library ‘adehabitatHR’ (Calenge 2006). We extracted mean snow cover values from each MCP from every monthly snow cover raster separately, using the R library ‘raster’ (Hijmans & van Etten 2023).

All calculations were performed using R version 4.2.2 ‘Innocent and Trusting’ (R Development Core Team 2022) and RStudio version 353 ‘Elsbeth Geranium’ (Posit team 2022).

Results

Foxtrot migration – always on the move

Earlier version: Except during the breeding season, Rough-legged buzzard migration continues throughout the year, even after the birds’ arrival at their traditionally recognized ‘wintering grounds’ (Figure 2a, b). For both quick and slow migrations, linear mixed-effects models with the season as a fixed factor received higher support from the likelihood ratio test (p < 0.001, Tables S1-S3). Rough-legged buzzards started their fall migration on 28 September (hereafter mean±sd for the day of the year: 271±11, n=31) and ended on 12 October (285±11, n=33). The mean latitude/longitude where the birds ended their fall migration was 55.57±1.92°/49.35±5.63°. During the winter, birds continued to migrate at a slower pace down to 49.53±2.01° latitude (on 5 February, 36±40, n=23) and 34.29±5.11° longitude (on 24 January, 24±47, n=23). Afterward, during the second phase of the slow winter migration, the birds returned to 55.52±2.63° latitude and 49.79±8.24° longitude to start the spring migration. Rough-legged buzzards started their spring migration to the Arctic on 27 April (117±7, n=27) and arrived at the breeding grounds on 15 May (135±8, n=18).

Revised version: Except during the breeding season, Rough-legged buzzard migration continues throughout the year, even after the bird’ arrival at their traditionally recognized ‘wintering ground’ (Figure 2a, b). For both quick and slow phases of the foxtrot migration, linear mixed-effects models with the season as a fixed factor received higher support from the likelihood ratio test (p < 0.001, Tables S1-S3). Rough-legged buzzards started their fall migration (quick phase) on 28 September (hereafter mean±sd for the day of the year: 271±11, n=31) and ended on 12 October (285±11, n=33). The mean latitude/longitude where the birds ended their fall migration was 55.57±1.92°/49.35±5.63° (Figure 2a). During the winter, birds continued to migrate at a slower pace down to 49.53±2.01° latitude (on 5 February, 36±40, n=23) and 34.29±5.11° longitude (on 24 January, 24±47, n=23). Afterward, during the second part of the slow phase, the birds returned to 55.52±2.63° latitude and 49.79±8.24° longitude to start the spring migration (Figure 2a). Rough-legged buzzards started their spring migration to the Arctic on 27 April (117±7, n=27) and arrived at the breeding grounds on 15 May (135±8, n=18).

Added in the revised version: Foxtrot migration of Rough-legged buzzards.

Removed in the revised version: Quick and slow migration patterns of Rough-legged buzzards throughout an annual cycle, termed ‘foxtrot migration’ for quick referencing.

Earlier version: Q – quick migration. Qf – Quick fall migration (orange), Qs – Quick spring migration (yellow), S1 – Slow migration, 1st phase (light blue), S2 – Slow migration, 2nd phase (dark blue). a) Change in the latitude of 43 Rough-legged buzzards during the year, red line – mean latitude of all birds, black vertical lines – mean dates of start and end of the migration phases, blue vertical line – mean date of the minimum latitude. Grey, sky blue, and piggy pink shaded areas – standard deviation of the means. b) Migration map. c) Difference in the migration parameters between the migration periods. Lines on the direction plot (down, right) represent the mean value for each bird; arrows represent the mean direction.

Revised version: Q – quick phase. Qf – Quick fall phase (orange), Qs – Quick spring phase (yellow), S1 – Slow phase, 1st part (light blue), S2 – Slow phase, 2nd part (dark blue). a) Change in the latitude of 43 Rough-legged buzzards during the year, red line – mean latitude of all birds, black vertical lines – mean dates of start and end of the migration phases, blue vertical line – mean date of the minimum latitude. Grey, sky blue, and piggy pink shaded areas – standard deviation of the means. b) Migration map. c) Difference in the migration parameters between the migration phases. Lines on the direction plot (down, right) represent the mean value for each bird; arrows represent the mean direction.

Earlier version: Quick-slow migration features comparison

Revised version: Quick-slow phase features comparison

Earlier version: During quick migration, individual birds flew greater distances in a shorter time, i.e., at a faster rate, than during slow migration. After arriving at what is traditionally known as the wintering grounds, the direction of migration changed, so the direction of quick and slow migration also differed (Table 1, Figure 2c). The quick migration was 1415±50 km long, whereas the slow migration (one phase) was 1026±55 km, i.e., 389±60 km shorter (p<0.001, Table S4, Figure 2c). During quick migration, birds flew for 15±3 days, and one phase of slow migration lasted 100±4 days, i.e., 85±5 days longer (p<0.001, Table S5, Figure 2c). At the same time, the second phase of slow migration was 54±7 days shorter than the first (p<0.001, Table S5, Figure 2c). The migration speed was 104±6 km/day during the quick phase and 12±7 km/day during the slow phase, i.e., about eight times higher (p<0.001, Table S6, Figure 2c). During the fall migration, birds moved in the SSW direction (7±2 deg), then turned 50±3 deg (p<0.001, Table S7, Figure 2c) to the west and started their 1st slow phase until mid-winter. After that, they turned back to the NEE direction (57±2 deg) and performed their 2nd slow migration phase for several months until they turned 54±3 deg (p<0.001, Table S7, Figure 2c) to the north and started their spring migration. As a result of additional analysis of the effect of sex on migration length, we found no significant difference between the migration distances of males and females (Table S8).

Revised version: During quick phase, individual birds flew greater distances in a shorter time, i.e., at a faster rate, than during slow phase. After arriving at what is traditionally known as the wintering grounds, the direction of migration changed, so the direction of quick and slow phases also differed (Table 1, Figure 2c). The quick phase was 1415±50 km long, whereas the slow phase (one part) was 1026±55 km, i.e., 389±60 km shorter (p<0.001, Table S4, Figure 2c). During quick phase, birds flew for 15±3 days, and one part of slow phase lasted 100±4 days, i.e., 85±5 days longer (p<0.001, Table S5, Figure 2c). At the same time, the second part of slow phase was 54±7 days shorter than the first (p<0.001, Table S5, Figure 2c). The migration speed was 104±6 km/day during the quick phase and 12±7 km/day during the slow phase, i.e., about eight times higher (p<0.001, Table S6, Figure 2c). During the fall migration, birds moved in the SSW direction (7±2 deg), then turned 50±3 deg (p<0.001, Table S7, Figure 2c) to the west and started their 1st slow phase until mid-winter. After that, they turned back to the NEE direction (57±2 deg) and performed their 2nd slow phase for several months until they turned 54±3 deg (p<0.001, Table S7, Figure 2c) to the north and started their spring migration. As a result of additional analysis of the effect of sex on migration length, we found no significant difference between the migration distances of males and females (Table S8).

Vegetation land cover during migration

Removed in the revised version: During fast migrations, Rough-legged buzzards cross the forest zone, while during slow migrations, they migrate within the grassland and cropland zone (Figure 3). Rough-legged buzzards migrated fast across the tundra zone on the north in the Arctic and then through the

Removed in the revised version: Vegetation land cover during quick and slow migration.

Removed in the revised version: a) Quick migration. Spring and fall migration together. b) Slow migration. 1st and 2nd phases of the slow migrations together. c) Migration map.

Earlier version: taiga zone. Therefore, during quick migration, the three most common vegetation land cover types were forest (44.5±2.9 %, hereafter, percentage of all mean daily positions during migration annotated with the given vegetation type), shrublands (29.9±3 %), and grasslands (24.7±2.8 %, Figure 3a). During slow migration, the three most common vegetation land cover types were grasslands (65.1±6.2 %), croplands (26.9±6.3 %), and forests (4.9±1.4 %, Figure 3b). According to the linear mixed-effects models, the percentage of all vegetation land cover types differed between the slow and quick migration periods (p<0.001, Table S9), except for the urban lands. Urban lands were more common during the slow than quick migration (Figure 3). However, this type has been annotated for too few birds to make an adequate comparison.

Revised version: During quick phases, Rough-legged buzzards cross the forest zone, while during slow phase, they migrate within the grassland and cropland zone (Figure 3). Rough-legged buzzards migrated fast across the tundra zone on the north in the Arctic and then through the taiga zone. Therefore, during quick phase, the three most common vegetation land cover types were forest (44.5±2.9 %, hereafter, percentage of all mean daily positions annotated with the given vegetation type), shrublands (29.9±3 %), and grasslands (24.7±2.8 %, Figure 3a). During slow phase, the three most common vegetation land cover types were grasslands (65.1±6.2 %), croplands (26.9±6.3 %), and forests (4.9±1.4 %, Figure 3b). According to the linear mixed-effects models, the percentage of all vegetation land cover types differed between the slow and quick phases (p<0.001, Table S9), except for the urban lands. Urban lands were more common during the slow than quick phase (Figure 3). However, this type has been annotated for too few birds to make an adequate comparison.

Added in the revised version: Vegetation land cover during quick and slow phases.

Added in the revised version: a) Quick phase (spring and fall periods together). b) Slow phase (1st and 2nd parts together). c) Migration map.

Snow cover – the main reason for the dynamic winter range

Earlier version: During the slow winter migration, Rough-legged buzzards experienced snow cover ranging from 4.8±1.0% in October to 85.2±4.6% in February (Figure 4). If birds spent the winter in the place where they arrived after the fall migration, they would experience snow cover conditions ranging from 4.6±0.6% in October to 99.5±0.1% in February (Figure 4). And if birds fly directly to the southeast and stay there for the whole winter, they would experience snow cover conditions ranging from 1.4±0.2 % in October to 81.1±5.0 % in January (Figure 4). Thus, if birds fly immediately to the southwest and stay there until the end of the winter, they will find conditions with less snow cover in spring (p<0.001, Table S10). And if birds stay where they ended the fall migration, they will find themselves in situations with more snow cover (p<0.001, Table S10). In the latter case, the difference between real and hypothetical situations is not as pronounced (85.2% vs. 99.5%), but more means that snow cover will be close to 100% for several months in this hypothetical situation (Figure 4b).

Revised version: During the slow phase of the foxtrot migration, Rough-legged buzzards experienced snow cover ranging from 4.8±1.0% in October to 85.2±4.6% in February (Figure 4). If birds spent the winter in the place where they arrived after the fall migration, they would experience snow cover conditions ranging from 4.6±0.6% in October to 99.5±0.1% in February (Figure 4). And if birds fly directly to the southeast and stay there for the whole winter, they would experience snow cover conditions ranging from 1.4±0.2 % in October to 81.1±5.0 % in January (Figure 4). Thus, if birds fly immediately to the southwest and stay there until the end of the winter, they will find conditions with less snow cover in spring (p<0.001, Table S10). And if birds stay where they ended the fall migration, they will find themselves in situations with more snow cover (p<0.001, Table S10). In the latter case, the difference between real and hypothetical situations is not as pronounced (85.2% vs. 99.5%), but more means that snow cover will be close to 100% for several months in this hypothetical situation (Figure 4b).

Earlier version: Slow migration and snow cover conditions.

Revised version: Snow cover conditions during the slow phase of the foxtrot migration.

a) 95% Minimum convex polygons (MCPs) of Rough-legged buzzards during winter. Arrows indicate the direction of the movement across months. OCT – October, NOV – November, DEC – December, JAN – January, FEB – February, MAR – March, APR – April. b) Snow cover conditions for the real situation (black) and two hypothetical situations – if birds spend the winter in the place where they arrived after the fall migration (green) and if birds fly directly to the southwest and stay there all winter (red). Dots represent mean values, error lines – standard errors.

Added in the revised version: In some raptor species, adult females disperse further than males (Mearns & Newton 1984; Serrano et al. 2001; Bildstein 2006; Whitfield et al. 2009). While our analysis does not support this (Table S8), it should be acknowledged that our dataset comprises 35 females, eight males, and no juveniles. The unbalanced sex ratio and absence of immature individuals in our sample may have potential influences towards the observed movement patterns in this study.

Discussion

Added in the revised version: Our study confirmed the existence of a previously hypothesized bird migration pattern (García & Arroyo 1998) characterized by alternating quick and slow phases (Figure 1). Following the breeding season, birds quickly traverse unfavorable habitats and proceed slowly through winter in favorable environments influenced by external factors. Subsequently, they quickly cross unfavorable terrain once more to reach the breeding region in spring (Figure 2). This pattern, termed ‘foxtrot migration’, distinguishes itself from ‘itinerancy’ by featuring directional and continuous movements. Additionally, it differs from ‘forage-on-fly migration’ as it occurs during the non-breeding period. We propose referring to the non-breeding area as ‘dynamic area’, reflecting its temporal heterogeneity, where the region with the highest bird density continually changes throughout the non-breeding period.

Added in the revised version: Our study affirmed the presence of foxtrot migration and a dynamic non-breeding range in Rough-legged Buzzards. The taiga zone posed as unfavorable habitat during the quick phase of foxtrot migration (Figure 3), given the difficulty for Rough-legged Buzzards to locate open areas for hunting. Conversely, the grassland and cropland zone served as favorable habitats during the slow phase throughout the entire non-breeding period (Figure 3), offering numerous open areas for hunting. Snow cover was the external factor driving their continual 1000 km southwest movement during winter (Figure 4). Our analysis revealed that if Rough-legged Buzzards remained at their fall migration endpoint without moving southwest, they would encounter 14.4% more snow cover (99.5% vs. 85.1%, Figure 4). Although this difference may seem small (14.4%), it holds significance for rodent-hunting birds, distinguishing between complete and patchy snow cover. Simultaneously, if Rough-legged Buzzards immediately flew to the southwest and stayed there throughout winter, they would experience 25.7% less snow cover (57.3% vs. 31.6%, Figure 4). Despite a greater difference than in the first case, it doesn’t compel them to adopt this strategy, as it represents the difference between various degrees of landscape openness from snow cover. This indicates ecological tipping points in the environment that are not predicted by a slight linear increase in effect size.

Added in the revised version: This pattern is likely to be observed in many migratory species with distinct seasonal cycles in their non-breeding range. In our study, we focused on an Arctic species that resides in the mid-latitudes during the non-breeding period. In general, foxtrot migration is likely to be a common feature among species where snow cover significantly influences food availability, especially those specializing in rodents. Originally proposed for birds in Africa, where seasonal changes regulate food availability, it’s likely that many insectivorous birds that spend the non-breeding season in Africa also exhibit this form of migration. This type of migration and dynamic non-breeding range is widespread both geographically and taxonomically, with obvious conservation implications.

Added in the revised version: The conservation implications

Added in the revised version: The conservation implications of our study are twofold: 1) the use of mid-winter (Christmas) bird surveys to determine non-breeding range may yield inaccurate results for species with foxtrot migrations (Figure 5a), and 2) declines in abundance within a particular segment of the non-breeding range may indicate changes in range dynamics rather than widespread declines in species abundance (Figure. 5b).

Added in the revised version: The use of foxtrot migration research for the assessment of the non-breeding range and conservation status of a particular species.

Added in the revised version: a) Christmas counts, often used to determine winter range, may give a misleading representation of non-breeding range for species with foxtrot migration. b) Changing the conservation status of a species based on reduced abundance in a particular area may not be appropriate for species with foxtrot migrations. Red – breeding range, blue – winter range. Latin numbers indicate the month the birds are present in the area.

Added in the revised version: Christmas Bird Counts provide insight into the mid-winter distribution of a species. For example, based on our study of the Rough-legged Buzzard, the species is predominantly present in the southwestern portion of its non-breeding range during this period, with only a small proportion present in the rest of the region. As a result, a map of the non-breeding range may inaccurately depict the species as common in the southwest and extremely rare in the northeast. This map would be inaccurate because, during the entire non-breeding period, Rough-legged Buzzards spend both fall and spring in the northeastern part and only mid-winter in the southwesternmost part.

Added in the revised version: To address this, continuous year-round GPS tracking of the species provides a means to track bird locations throughout the non-breeding season, facilitating the creation of accurate distribution maps, particularly for species that exhibit foxtrot migrations and dynamic ranges. We advocate representing temporal heterogeneity (range dynamics) on maps as distinct zones, denoting periods when the species is abundant in a given area. To distinguish temporal from spatial heterogeneity, we recommend using lines to delineate the boundaries of these zones, rather than color shading, and incorporating numbers to denote the months of species abundance in a given zone (Figure 5a). We suggest ecologists include dynamic non-breeding ranges in descriptions and range maps for foxtrot migratory bird species.

Added in the revised version: Population counts for a species are often limited to a portion of its range. Therefore, conclusions about conservation status drawn from such counts may be misleading. A decline in abundance within a particular portion of the non-breeding range may indicate changes in range dynamics rather than a general decline in the species. For example, climate change may affect snow cover dynamics, reducing its intensity in northern regions. As a result, species whose range dynamics depend on snow cover may choose to remain in the northern areas and not migrate as far south as they traditionally have. Despite this shift in range dynamics, overall species abundance may remain unchanged. Similar patterns have affected Rough-legged Buzzards in some areas of the European non-breeding range.

Removed in the revised version: Our study has revealed a novel bird migration pattern that likely occurs in many migratory species inhabiting winter areas with pronounced seasonal cycles. This pattern results in the dynamic shift of the birds’ non-breeding range by hundreds of kilometers. In contrast to traditional migration patterns where birds settle in a single location after the initial rapid migration phase in the fall (Bildstein 2006; Newton 2008), the quick-and-slow migration involves a continuous migration at a slower pace through the winter range, followed by a swift return to the breeding site (Figures 2a, b; 4a; 5a). This pattern features a consistent, directional return migration throughout the winter, which is repeated annually, unlike the nomadic movements observed in some bird species during winter (Eng & Schladweiler 1972; Roth & Vetter 2008; Heckscher et al. 2011). We, therefore, termed this migration pattern ‘foxtrot’ migration – a quick movement followed by slow movements in different directions – for quick referencing. Also, unlike the ’fly-and-forage’ migration pattern (Strandberg & Alerstam 2007) and long stopovers during migrations (Shimada et al. 2014) caused by physiological needs, in this case, birds are moving primarily due to external factors that are changing rapidly in the present-day world as a result of climate change. Therefore, this pattern is potentially subject to significant change during the Anthropocene, a change that could have cascading effects on species distributions.

Earlier version: A 2022 Dutch study found a decline in wintering Rough-legged buzzards over the last 40 years (Hornman, Boele & van Winden 2022). On the one hand, this may represent a conservation concern. On the other hand, applying the rationale of the foxtrot migration, the apparent local decline may simply be attributed to climate change, resulting in less comprehensive snow coverage in the northeastern wintering areas of Rough-legged buzzards relative to the Netherlands. Such a shift in snow coverage makes it less probable for the birds to migrate to the Netherlands for overwintering. This proposition is further supported by a study of the winter population dynamics of Rough-legged buzzards in the Netherlands in 2011, showing that the main winter population peak occurred in late December, with many birds migrating (Vansteelant, Faveyts & Buckens 2011). The authors also found that the main migration occurred after heavy snowfall in northern Europe, supporting our ‘foxtrot’ migration hypothesis. We suggest that population changes of highly mobile organisms in many non-breeding areas around the world could be similarly explained by a more or less continuous movement of individuals in response to climate alterations.

Revised version: A 2022 Dutch study found a decline in wintering Rough-legged buzzards over the last 40 years (Hornman, Boele & van Winden 2022). On the one hand, this may represent a conservation concern. On the other hand, applying the rationale of the foxtrot migration, the apparent local decline may simply be attributed to climate change, resulting in less comprehensive snow coverage in the northeastern wintering areas of Rough-legged buzzards relative to the Netherlands. Such a shift in snow coverage makes it less probable for the birds to migrate to the Netherlands for overwintering. This proposition is further supported by a study of the winter population dynamics of Rough-legged buzzards in the Netherlands in 2011, showing that the main winter population peak occurred in late December, with many birds migrating (Vansteelant, Faveyts & Buckens 2011). The authors also found that the main migration occurred after heavy snowfall in northern Europe, supporting our foxtrot migration explanation for this decline. Therefore, investigating the dynamic range and foxtrot migration is critical to understanding a species’ range and effectively assessing its conservation status.

Added in the revised version: Significant environmental changes, including changes in snowpack dynamics, have occurred in recent decades that may have significant impacts on bird migration (Tucker et al. 2018; Sumasgutner et al. 2021). In order to understand the effects of global environmental change on bird migration behavior, there is an urgent need for in-depth investigations into the intricate relationship between migration patterns and evolving environmental factors. The advent of new technologies in recent decades has expanded our understanding of bird migration, enabling the identification of novel behavioral patterns and facilitating the study of their susceptibility to various environmental influences. This knowledge is becoming increasingly urgent in the face of rapid climate change and anthropogenic habitat alteration, which pose threats to both wildlife and human populations.

Removed in the revised version: Understanding the dynamic range and migration pattern of Rough-legged buzzards is crucial for developing effective conservation strategies for this, but also for other species. In conservation planning, balancing prohibited management activities within the species’ range and the benefits of such activities is a key consideration (Beatley 1992; Jones, Pejchar & Kiesecker 2015; Harfoot et al. 2018; Carrell, Hammill & Edwards 2022). By understanding the dynamic range of the species, a more optimal trade-off can be identified, allowing management activities to be carried out when the species is not present. We recommend that researchers working on species distribution atlases and range maps take this important aspect of a species’ biology into account. In most of the bird atlases and web resources of today (Keller et al. 2020; Audubon 2022), the winter range is usually depicted as a single entity or an area where the species may be present in varying densities (Figure 5b, c), indicating spatial heterogeneity. We propose that temporal heterogeneity be identified as separate zones, indicating periods when the species is abundant, perhaps beyond a threshold value, in the area. To differentiate from spatial heterogeneity, we suggest using lines to define the boundaries of these zones instead of using color shading (Figure 5c) or arrows to show the migration phases (Figure 5a). It is important to note that when mid-winter counts (e.g., Christmas counts) are used to determine the winter range, the spatial heterogeneity observed may be due to temporal heterogeneity. This study indicates that the fact that the majority of birds are far from the northern limit of their winter range in mid-winter does not necessarily mean that they are less abundant there throughout the winter compared to other parts of their winter range.

Removed in the revised version: Winter range representations.

Removed in the revised version: Red – breeding range, blue – winter range. a) Migration cycle of Rough-legged buzzard; arrows represent migration (solid line – quick migration, dashed line – slow migration). b) Winter range without heterogeneity shown. c) Winter range with spatial heterogeneity. d) Winter range with temporal heterogeneity; the numbers indicate the months the birds are common in the outlined area.

Removed in the revised version: While the practical importance of this migratory behavior is clear, understanding the reasons behind this pattern of migration is crucial for predicting potential changes in response to climate change. As in the case of species wintering in subtropical or tropical parts of Africa, the movement patterns are likely driven by the interplay between seasonal cycles and habitat types – similar to what we propose for the alternating periods of fast and slow migration in Rough-legged buzzards. During the quick migration phase, birds pass through habitats that are unfavorable for their survival, while during the slow migration phase, they navigate through more suitable habitats. In our study, the quick migration tracks, as shown in Figures 2 a, b and 3, traverse the taiga zone, which is not conducive to the hunting needs of Rough-legged buzzards, which require open areas for hunting (Ferguson-Lees & Christie 2001; Bechard & Swem 2002). Upon arriving in an appropriate habitat, such as a wooded field and croplands, the birds temporarily halt or slow their migration but then turn from the south to the southwest and gradually migrate in this direction as the winter progresses (Figures 2c and 4). The question arises as to why the birds turn to the southwest and do not continue their migration further south into the steppe zone, which is an open habitat, too. We assume that the lack of forests in the steppe, which serve as crucial roosting sites for Rough-legged buzzards, prevents them from doing so. Our analysis of the locations assigned to forests during the slow migration revealed a scarcity of such sites, with less than 5% of the locations being assigned to forests (Figure 3b).

Removed in the revised version: Despite their scarcity in these habitats, forests appear to hold significant importance for Rough-legged buzzards for nocturnal safety. Another question is why Rough-legged buzzards do not stay in places where their autumn migration has ended. Based on our analysis, we postulate that the primary driver behind the winter migration of Rough-legged buzzards is the southwestward progression of the snow line over the winter season.

Removed in the revised version: The progression of snow cover in the wintering range of Rough-legged buzzards plays a significant role in their winter migration pattern. The winter range of Rough-legged Buzzards experiences a dynamic snow cover with a gradual increase from October to January-February and a subsequent gradual decrease from February to May, moving from the northeastern part of the range to the southwestern part. Correspondingly, the buzzards also move from northeast to southwest during the winter, as if in response to the increasing snow (Figures 2b, 4a). They only partially avoid snow, with snow cover reaching 85% in the southeast of their range by mid-winter. However, they successfully avoid continuous snow cover, with snow cover reaching 100% in the northeast part of their range (Figure 4b). There is a critical limit of snow cover, around 90%, beyond which hunting becomes infeasible. The shortening day length in the northeast compared to the southwest and the fact that Rough-legged buzzards only hunt during the day (Pokrovsky et al. 2021) highlights the significance of even a tiny difference in snow cover. One might wonder why the buzzards would not move to the southwest of the range in the autumn and remain there throughout the winter. This behavior might be attributed to optimizing migration energy costs, with the buzzards moving to the optimal habitat only when the prevailing conditions become unfavorable for their survival. This is in accordance with the theory of optimal migration (Alerstam & Lindström 1990; Houston 1998; Newton 2008). However, some individuals may migrate directly to the southwest, reflecting the plasticity of the migration behavior (Figure S1). The dynamics of snow cover hold a significant influence not only over the winter migration patterns but also in the entire migration cycle of Rough-legged buzzards.

Removed in the revised version: Both quick and slow migration phases in Rough-legged buzzards are influenced by environmental changes that are part of seasonal cycles (Sonerud 1986; Curk et al. 2020). These cycles result in variations in temperature, precipitation, snow cover, and day length, which affect food availability. Latitude affects the extent of these seasonal changes. In autumn, birds in northern breeding grounds face declining temperatures, shorter days, and increased snow cover, prompting a migration to escape these adverse conditions (Alerstam 1991; Houston 1998; Alerstam, Hedenström & Åkesson 2003). However, unfavorable habitats in the southern portion of their breeding range (taiga zone) prevent gradual movement away from adverse environmental conditions. Hence, they perform quick migration to cross the forest zone. Once in favorable habitats, the birds are temporarily insulated from the negative factors they had previously fled, allowing for respite. As the influence of these negative factors eventually reaches the birds’ favorable habitats in the south, they must migrate again, but at a slower pace that corresponds with the gradual approach of these factors. This delay results in a slower migration rate in the first half of winter compared to the second half (Figure 2, Table S2).

Removed in the revised version: In some raptor species, adult females disperse further than males (Mearns & Newton 1984; Serrano et al. 2001; Bildstein 2006; Whitfield et al. 2009). While our analysis does not support this (Table S8), it should be acknowledged that our dataset comprises 35 females, eight males, and no juveniles. The unbalanced sex ratio and absence of immature individuals in our sample may have potential influences towards the observed movement patterns in this study.

Removed in the revised version: The quantification of environmental changes that could prove fatal to bird species presents yet another challenge for conservation efforts in an era of rapid global change. In this study, we have demonstrated the impact of seemingly minor variations in ecological parameters on the species’ behavior. For instance, the presence of small patches of woodland in the winter range might appear crucial to the survival of the Rough-legged buzzard. Elimination of these seemingly minor elements of vegetation cover through management actions could have dire consequences for the species. Similarly, snow cover, while not necessarily a limiting factor in the distribution of a particular species, can become critical at certain thresholds and play a crucial role in shaping species’ behavior. As our study shows, if Rough-legged buzzards fly southwest immediately and stay there for the whole winter, they will experience 25.7% (57.3% vs. 31.6%) less snow cover in March, but this does not force them to adopt this strategy. At the same time, if Rough-legged buzzards stay where they finished their fall migration and don’t move southwest, they will experience 14.4% (99.5% vs. 85.1%) more snow cover in February, and this will force them to migrate. The point is that in this case, although the difference is not as great (14.4% vs. 25.7%), it is the difference between complete snow cover and inconsistent snow cover. This again shows that when we discuss conservation strategies, it is not necessarily important to judge such parameters in terms of numbers or percentages, but rather to understand the context of these parameters in the species’ biology and their importance for the species’ existence. In some contexts, significant environmental changes may not affect the species’ life, but in other contexts, even small changes in the same environmental parameters may be fatal to the species: in short, there may be tipping points in environmental settings that are not predicted by a small linear increase in effect sizes.

Conclusions

Added in the revised version: Our study has identified and characterized a new pattern of migratory behavior, the ‘foxtrot migration’, along with the associated concept of ‘dynamic range’. This discovery has significant implications for conservation strategies and adequate representation of non-breeding habitats. As animal tracking technology advances, we expect to discover more new details about animal migrations and, consequently, new patterns of migratory behavior. With these technological developments, we are gaining a clearer understanding of the complex patterns and behaviors that drive animal migrations, providing valuable insights for conservation efforts and deepening our understanding of the natural world.

Removed in the revised version: The quick-and-slow migration pattern of Rough-legged buzzards (we termed them ‘foxtrot migration’) results in a significant temporal shift of their winter range of over 1000 km. After the birds arrive in their ‘wintering range’, they move slowly but continuously towards the southwest and back during the winter, highlighting the adaptability of birds in response to changing environments. Our findings have crucial implications for conservation efforts in the Anthropocene, where rapidly changing environmental factors can severely impact bird migration (Tucker et al. 2018; Sumasgutner et al. 2021). We suggest that ecologists should include dynamic winter ranges in species descriptions and range maps to develop effective conservation strategies for the bird species. The knowledge about dynamic range shifts is imperative in the face of rapid climate and human-induced habitat changes threatening wildlife and humans.

Supporting information